<center>
<img src="logo.png" height="900"> 
</center>


# Максимально правдоподобная домашка

В этом задании мы с вами отточим метод максимального правдоподобия.

```python
import numpy as np
import pandas as pd

from scipy import stats
from scipy.optimize import minimize

import seaborn as sns
import matplotlib.pyplot as plt

plt.style.use('ggplot')
%matplotlib inline
```

## Упражнение 1 (Разминка)

__а)__ Пусть $X \sim Exp(0.05)$. Сгенерируйте из этого распределения выборку размера $1000$. У метода `rvs` используйте `random_state=42` 

Обратите внимание, что `stats.expon` принимает в качестве параметра `scale` значение $\frac{1}{\lambda}$, а $0.05$ это именно $\lambda$.

<!-- grade_id: cell-f56ad0d7e5287be8, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

x = ...

# your code here

```

```python
assert np.abs(19.45 - np.mean(x)) < 0.01
```

__б)__ Выпишите логарифмическую функцию правдоподобия на бумажке. Затем вбейте её в python.

<!-- grade_id: cell-8b65f97b3713651c, points: 0, grade: false, solution: true -->
```python
def lnL(theta, x):
    
    # your code here
    
    return -l
```

<!-- grade_id: cell-06138bffb4961739, points: 1, grade: true, solution: false -->
```python
assert np.abs(lnL(4, x) - 76414.1814) < 1e-4

# несколько похожих скрытых тестов
```

__в)__ Оцените значение параметра $\lambda$ по сгенерированной выборке `x`. Аккуратнее выбирайте точку для инициализации параметра. 

<!-- grade_id: cell-46d606a0b318889e, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

theta_init = ...
res = ...

# your code here

```

<!-- grade_id: cell-fdb20543b8303077, points: 1, grade: true, solution: false -->
```python
assert np.abs(res.x[0] - 0.0514) < 1e-4

# несколько похожих скрытых тестов
```

__г)__ Постройте график для функции правдоподобия. По оси $x$ отложите значения $\lambda$, по оси $y$ логарифм правдоподобия. Отметьте на картинке точку оптимума. 

<!-- grade_id: cell-4de953f7f7c66a39, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

__д)__ Постройте для параметра $\lambda$ асимптотический $80\%$ доверительный интервал. 

<!-- grade_id: cell-c20e8fc746ce9925, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

alpha = ...
z = ...

left = ...
right = ...

# your code here

```

<!-- grade_id: cell-b6d9034e824cb106, points: 2, grade: true, solution: false -->
```python
assert np.abs(z - 1.2815) < 1e-4

# несколько похожих скрытых тестов
```

__е)__ Проверите на уровне значимости $1\%$ тестом отношения правдоподобия гипотезу о том, что $\lambda = 0.05$. 

<!-- grade_id: cell-78a916d342e6ca21, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

lnL_un = ...
lnL_r = ...

# наблюдаемое значение статистики
LR_obs =  ... 

# критическое значение
LR_cr = ...

# your code here

```

<!-- grade_id: cell-f2b0554f56b04198, points: 1, grade: true, solution: false -->
```python
assert np.abs(lnL_un - 3967.8531) < 1e-4

# несколько похожих скрытых тестов
```

Справились? Значит вас ничто не может остановить реализовать подобный пайплайн для более сложной ситуации. Убедитесь в этом сами, решив задачку про «Челленджер».

## Упражнение 2 (Логистическая регрессия)

В этом упражнении мы оценим методом максимального правдоподобия логистическую регрессию.

$28$ января $1986$ года космический шатл «Челленджер» взорвался при взлёте. Семь астронавтов, находившихся на борту, погибли. В ходе расследования причин катастрофы основной версией была неполадка с резиновыми уплотнительными кольцами в соединении с ракетными ускорителями. Для $23$ предшествовавших катастрофе полётов «Челленджера» известны температура воздуха и появление повреждений хотя бы у одного из уплотнительных колец.

```python
df = pd.read_csv('challenger.csv', sep=',')
print(df.shape)
df.head()
```

__а)__ Постройте на графиках гистограмму для распределения температуры $(Temperature)$ в ситуациях, когда хотябы одно кольцо было повреждено $(Incident = 1)$ и в ситуациях, когда ни одного повреждения не было $(Incident = 0)$. 

<!-- grade_id: cell-4f7c85bfda7decc0, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

```python
df[df['Incident'] == 1].Temperature.mean(), df[df['Incident'] == 0].Temperature.mean()
```

__б)__ Судя по всему, между температурой и вероятностью поломки есть взаимосвязь. Пусть вероятность поломки зависит от температуры экспоненциально:
    
$$
P(Incident = 1) = \frac{e^{\alpha + \beta \cdot t}}{1 + e^{\alpha + \beta \cdot t}} = \frac{1}{1 + e^{-1\cdot(\alpha + \beta \cdot t)}} 
$$
    
Выведите функцию правдоподобия и задачу для оптимизации поиска коэффициентов $\alpha$ и  $\beta$ на бумажке, а затем вбейте её в python. 

<!-- grade_id: cell-0b993a930584fc58, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

def lnL(theta, data):
    a, b = theta
    
    # your code here
    
```

<!-- grade_id: cell-63d24d1dbe5d8980, points: 1, grade: true, solution: false -->
```python
assert np.abs(lnL([4, 0.2], df) - 294.8) < 1e-4

# несколько похожих скрытых тестов
```

__в)__ Оцените параметры $\alpha$ и $\beta$ с помощью метода максимального правдоподобия.

<!-- grade_id: cell-160598072dae63da, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

theta_init = ...
res = ...

# your code here

```

<!-- grade_id: cell-17af0e0290ffd6bf, points: 1, grade: true, solution: false -->
```python
assert np.abs(res.fun - 10.1576) < 1e-4

# несколько похожих скрытых тестов
```

__г)__ С помощью теста отношения правдоподобий на уровне значимости $5\%$ проверьте гипотезу о том, что $\alpha = \beta = 0$.

<!-- grade_id: cell-d58c71fd7ff43669, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

lnL_un = ...
lnL_r = ...

LR_obs = ...
LR_crit = ...

# your code here

```

<!-- grade_id: cell-1144bb1c976e84e0, points: 1, grade: true, solution: false -->
```python
assert np.abs(lnL_un - 10.1576) < 1e-4

# несколько похожих скрытых тестов
```

__д)__ В день катастрофы температура уплотнителя была $31$ градус по Фарингейту. Предскажите с помощью построенной модели вероятность катастрофы (для этого подставьте в формулу поиска вероятности получившиеся оценки и температуру).

<!-- grade_id: cell-dc73e72855a06d26, points: 0, grade: false, solution: true -->
```python
xnew = 31
alpha, beta = res.x

### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

p_hat = ...

# your code here

```

<!-- grade_id: cell-7fbcaf3d5937cf8c, points: 1, grade: true, solution: false -->
```python
assert np.abs(p_hat - 0.9996) < 1e-4

# несколько похожих скрытых тестов
```

Только что мы с вами оценили логистическую регрессию. На самом деле функцию потерь, которую обычно используют для оценки логистической регрессии (logloss) можно получить из метода максимального правдоподобия. Об этом мы поговорим на следующей неделе нашего курса. 

__е)__ А сейчас, давайте сделаем финальный штрих и оценим ту же самую моделю с помощью пакета `statsmodels`.

```python
import statsmodels.formula.api as smf

model = smf.logit(data=df, formula='Incident ~ 1 + Temperature').fit()
model.summary()
```

Посмотрев на протокол, получившийся в результате оценивания модели ответьте на вопросы: 

- Чему равноо значение логарифма правдоподобия? Запишите это значение в `logLik`
- Чему равна оценка константы $\alpha$? Запишите это значение в `alph`
- По свойствам метода максимального правдоподобия $\hat{\beta} \sim N(\beta, \hat{\sigma}^2_{\beta})$, чему равно значение $\hat{\sigma}_{\beta}$? Запишите это значение в `std_beta`
- Чему равно p-значение для гипотезы $H_0: \beta = 0$? Запишите это значение в `pval_beta`
- Отвергается ли гипотеза $H_0: \beta = 0$ на $5\%$ уровне значимости? Если да, запишите в `is_h` значение `False`

<!-- grade_id: cell-47498286ac9be898, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

logLik = ...
alph = ...
std_beta = ...
pval_beta = ...
is_h = ...

# your code here

```

<!-- grade_id: cell-738394af91e6aa78, points: 1, grade: true, solution: false -->
```python
assert np.abs(logLik + 10.158) < 1e-4

# несколько похожих скрытых тестов
```