<center>
<img src="logo.png" height="900"> 
</center>


#  Аварии на шахтах

```python
import numpy as np

import scipy.stats as sts
import matplotlib.pyplot as plt

plt.style.use('ggplot')  # стиль для графиков
%matplotlib inline
```

На угольных шахтах ежегодно происходят аварии. Англия довольно давно собирает информацию о числе крупных аварий. Именно её нам предстоит проанализировать:

* хочется понять как часто в среднем происходят аварии 
* насколько большой у среднего числа аварий доверительный интервал

```python
import pandas as pd

df = pd.read_csv("coals.csv", index_col=0)
df.head()
```

В колонке `year` находится год наблюдений. В колонке `count` число аварий за этот год на шахтах в Англии. Посмотрим на распределение числа аварий. 

```python
df['count'].hist(bins=30);
```

Судя по гистограмме уместно предположить, что число аварий имеет распределение Пуассона, $Poiss(\lambda)$.  То есть, если мы хотим что-то понять про среднее число аварий, нам надо оценить параметр $\lambda$. 

__а) Метод моментов__

Найдите оценку параметра $\lambda$ с помощью метода моментов. Запишите получившийся результат в переменную `lambda_hat`.

<!-- grade_id: cell-a4673b56f9be221f, points: 0, grade: false, solution: true -->
```python
n_hat = df.shape[0]   # число наблюдений 
lam_hat = ...         # ваша оценка метода моментов

# your code here

```

<!-- grade_id: cell-3f10ac75755a9483, points: 1, grade: true, solution: false -->
```python
# проверка, что оценка была найдена корректно 
# похожие тесты, но для lam_hat скрыты от вас
# они запустятся после того как вы сдадите тетрадку на проверку в грейдер

assert n_hat == 112 

```

__б) Стандартное отклонение__ 

Среднее, оценённое методом моментов будет иметь асимптотически нормальное распределение. Оцените стандартное отклонение этого распределения по выборке и запишите ответ в переменную `lam_se`. 

> __Примечание:__ мы предполагаем, что выборка пришла из распределения Пуассона. Для него $Var(X) = \lambda$. Для оценки стандартного отклонения используйте этот факт. Конечно же вы можете оценить его через функцию `std`, в плане асимптотики результат получится одинаковый. Оба способа корректны. Но с точки зрения грейдера, у вас будет ошибка. 

<!-- grade_id: cell-d89f8ff08d240616, points: 0, grade: false, solution: true -->
```python
lam_se = ...         # ваша оценка

# your code here

```

<!-- grade_id: cell-69d840309379a00d, points: 1, grade: true, solution: false -->
```python
# проверка, что оценка была найдена корректно 
assert lam_se < 1.6474

# несколько похожих тестов скрыты от вас

```

__в) Доверительный интервал__ 

Постройте для неизвестного математического ожидания $\lambda$ доверительный интервал. Возьмите уровень значимости равный $10\%$. В переменные `left` и `right` запишите соотвествующие границы интервала. В переменную `delta` запишите его ширину.

<!-- grade_id: cell-729da6fec593102e, points: 0, grade: false, solution: true -->
```python
alpha = ...

left = ...
right = ...
delta = ...

# your code here

```

<!-- grade_id: cell-681b7ef69959cc8f, points: 1, grade: true, solution: false -->
```python
# проверка, что границы доверительного интервала найдены верно
assert np.abs(left - 1.502389916114217) < 1e-5

# несколько похожих тестов скрыты от вас

```

__г) Число катастроф__  

Многие говорят, что в $20$ веке катастроф стало на порядок меньше. Давайте проверим это. Постройте оценку для $\lambda_1$ по данным до $1900$ года включительно. Постройте оценку для $\lambda_2$ по данным после $1900$ года. Для разности $\lambda_1 - \lambda_2$ постройте $95\%$ доверительный интервал. 

<!-- grade_id: cell-a52534c295035fea, points: 0, grade: false, solution: true -->
```python
x = df[df.year <= 1900]['count']
y = df[df.year > 1900]['count']

alpha = ...

left = ...
right = ...
delta = ...

# your code here

```

<!-- grade_id: cell-40e69d1d970e89ea, points: 2, grade: true, solution: false -->
```python
# проверка, что границы доверительного интервала найдены верно
assert x.size == 50
assert np.abs(left - 1.2835473859716393) < 1e-5

# несколько похожих тестов скрыты от вас

```

Ноль не вошёл в этот доверительный интервал. Это означает, что на уровне значимости $5\%$ мы можем сказать, что число аварий правда снизилось. 

__д) Вероятность для распределения Пуассона__ 

Вспомните формулу, по которой считается вероятность того, что Пуассоновская случайная величина примет какое-то значение. Найдите точечную оценку вероятности того, что произойдёт только одна авария. В качестве $\lambda$ используйте оценку метода моментов, полученную в первом задании. 

<!-- grade_id: cell-1383c35748bfed27, points: 0, grade: false, solution: true -->
```python
def poiss_prob(k):
# your code here

```

<!-- grade_id: cell-301364eca6748674, points: 1, grade: true, solution: false -->
```python
# проверка,функции
assert np.abs(poiss_prob(1) - 0.3098761) < 1e-5

# похожие тесты скрыты от вас

```

__е) Дельта-метод__ 

Та вероятность, которую мы нашли в предыдущем пункте - это лишь точечная оценка. Она является функцией от среднего. С помощью дельта-метода постройте для этой вероятности $95\%$ довертельный интервал.

<!-- grade_id: cell-cd0a438c7adffcdb, points: 0, grade: false, solution: true -->
```python
alpha = ...

left = ...
right = ...
delta = ...

# your code here

```

<!-- grade_id: cell-94c00ca69c5dadef, points: 4, grade: true, solution: false -->
```python
# проверка, что интервал построен корректно
assert np.abs(left - 0.279) < 1e-3

# похожие тесты скрыты от вас

```

По аналогии можно построить доверительные интервалы для вероятности того, что случайная величина примет значения от нуля до девяти. Проделайте это в качестве необязательного упражнения и оформите процедуру расчёта в виде функции.

<!-- grade_id: cell-957ca214e0fc17cb, points: 0, grade: false, solution: true -->
```python
def prob_ci(k, alpha = 0.01, lam_hat=lam_hat):
# your code here

```

```python
left, right = prob_ci(1)
delta = right - left

assert np.abs(delta - 0.08147) < 1e-5

left, right = prob_ci(0)
delta = right - left

assert np.abs(delta - 0.11551) < 1e-5
```

Сделали функцию? Запускайте код для не очень информативной картинки!

```python
poiss_rv = sts.poisson(lam_hat)

x = np.arange(10)

p = poiss_rv.pmf(x)
left, right = prob_ci(x)

plt.figure(figsize=(10,5))
plt.plot(x, p, 'bo', ms=6, label='poisson pmf')
plt.fill_between(x, left, right, alpha=0.5)
plt.vlines(x, 0, p, lw=1)
plt.legend(loc='best', frameon=False)
plt.show()
```

```python
for k,(i,j) in enumerate(zip(left, right)):
    print(f"Вероятность P(X={k}) лежит между  {i:.5} и {j:.5}")
```

Интервалы получился довольно короткими. Попробуйте поиграться с уровнем значимости и посмотреть как они меняются в зависимости от него. 

В этой терадке нам осталось обсудить последний нюанс. Мы из теории вероятностей знаем, что для случайной величины $X$, имеющей распределение Пуассона $\mathbb{E}(X) = Var(X) = \lambda$. Для текущей выборки мы видим, что между дисперсией и средним есть небольшое отличие. 

```python
df['count'].mean()
```

```python
df['count'].var()
```

Не факт, что мы имеем право предполагать, что наша выборка пришла к нам из распределения Пуассона. Все предпосылки о распределениях выборки нужно тестировать. Позже мы научимся это делать.