<center>
<img src="logo.png" height="900"> 
</center>


# Линейная регрессия 

Данные для этой домашки взяты из [книги по эконометрике Филиппа Картаева](https://www.econ.msu.ru/departments/mmae/Category.20190227154534_1269/Category.20190227155254_6724/)

```python
import numpy as np
import pandas as pd

from scipy import stats 

import seaborn as sns
import matplotlib.pyplot as plt

plt.style.use('ggplot')
%matplotlib inline
```

# Задание 1 (тренинг)

Руководство крупной торговой сети планирует выяснить, помогает ли тренинг по продажам увеличить эффективность работы менеджеров по продажам. 

Для решения этой задачи были собраны данные: 

- `sales` - объём продаж данного менеджера (в тысячах рублей за период)
- `training` - дамми-переменная, равная единице, если в а самом начале данного периода менеджер прошел тренинг по продажам (работники, которые направлялись на курсы, выбирались из общей совокупности работников компании при помощи специальной лотереи)
- `female` - дамми-переменная, равная единице для менеджеров женщин и нулю для мужчин
- `experience` - опыт работы менеджера в годах 
- `capital` -  дамми-переменная, равная единице, если менеджер работает в столичном отделении компании, и равная нулю в противном случае
- `IQ` - все менеджеры при приеме на работу проходят IQ-тест 

```python
df = pd.read_csv('training.tsv', sep='\t')
print(df.shape)
df.head()
```

Предполагается, что в данных есть гетероскедастичность. Все гипотезы проверяйте с учётом коррекции `HC3`.

> Гетероскедастичность довольно часто присутствует в данных. Из-за этого, на практике, обычно, сразу вводят коррекцию на неё, и тестируют все гипотезы с помощью асимптотического подхода.

__[а]__ Оцените регрессию переменной `sales` на переменные `training`, `female`, `experience`, `capital`, `IQ`. Значимо ли уравнение в целом на уровне значимости 5%? Сколько переменных значимы на уровне значимости $5\%$? 

<!-- grade_id: cell-149c04bce064e935, points: 0, grade: false, solution: true -->
```python
import statsmodels.formula.api as smf

### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

- Запишите в переменную `prob_f` $p$-value для гипотезы о незначимости регрессии. 
- В переменную `cnt` запишите число незначимых коэффициентов. 
- На сколько в среднем, при прочих равных условиях, прохождение тренинга увеличивыет объём продаж менеджера? Запишите ответ в переменную `tren`.

<!-- grade_id: cell-f0be3ebdb49cf2b0, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

prob_f = ...
cnt = ...
tren = ...

# your code here

```

<!-- grade_id: cell-aa4d47500164f617, points: 1, grade: true, solution: false -->
```python
assert np.abs(prob_f - 7.03e-159) < 1e-4

# несколько похожих скрытых тестов
```

__[б]__ Оцените модель заново, исключив из неё переменную `capital` и добавив переменную `training:capital` (произведение столбцов). Значим ли коэффициент перед этой переменной на уровне значимости $5\%$? 

<!-- grade_id: cell-6fc9a46448b3ee92, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

- На сколько в среднем,  при прочих равных условиях, увеличивыет объём продаж столичного менеджера? Запишите ответ в переменную `ans1`.
- На сколько в среднем,  при прочих равных условиях, увеличивыет объём продаж менеджера не из столицы? Запишите ответ в переменную `ans2`.

<!-- grade_id: cell-225d2b4b1df7de9d, points: 0, grade: false, solution: true -->
```python
ans1 = ...
ans2 = ...

# your code here

```

<!-- grade_id: cell-73c14915270808fc, points: 1, grade: true, solution: false -->
```python
assert np.abs(ans2 - 15.7) < 1e-1

# несколько похожих скрытых тестов
```

__[в]__ Для модели, оценённой в пункте б) проведите тест Бреуша-Пагана на гетероскедастичность. Действительно ли она наблюдается ли она в остатках? 

<!-- grade_id: cell-41d5fed492ba5cd8, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

## Задание 2 (размер класса) 

Давайте попробуем проанализировать данные о том, как размер класса влияет на успеваемость учеников. Коррекцию HC3 в этом задании мы использовать не будем. 

__Описание переменных:__

- `CLASS` - размер класса, в котором обучается школьник, воздействие именно этой переменной нас интересует 
- `EXPN` - средние расходы на одного школьника в школе, где он учится, измеренные в тысячах долларов в год
- `INCOME` - средний доход на одного члена семьи в семие школьника, измеренный в тысячах долларов в год 
- `TEST` - результат итогового стандартизированного теста, который писали все школьники в конце учебного года, она характеризует качество обучения 

```python
df = pd.read_csv('study.tsv', sep='\t')
print(df.shape)
df.head()
```

```python
sns.pairplot(df);
```

__[a]__ Оцените параметры модели 

$$
TEST_i = \beta_0 + \beta_1 \cdot CLASS_i + \varepsilon_i
$$

Коррекцию на гетероскедастичность вводить не надо. 

<!-- grade_id: cell-d9033f154a5b7e01, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

Значима ли переменная `CLASS`? В качестве ответ на вопрос вбейте получившееся значение $p$-value и получившееся значение коэффициента.

<!-- grade_id: cell-fa8531f95cd582ab, points: 0, grade: false, solution: true -->
```python
prob = ...
beta = ...

# your code here

```

<!-- grade_id: cell-77bc8b71cbd2ff68, points: 1, grade: true, solution: false -->
```python
assert prob < 1e-10

# несколько похожих скрытых тестов
```

В получившемся уравнении коэффициент $\beta_1$ можно проинтерпретировать так, что  каждый дополнительным человек в классе приводит к тому, что результат типичного школьника из этого класса снижается в среднем на $beta_1$ балла.

> Мы можем так интерпретировать значение коэффициента, если выполняются предпослыки линейной регрессии. Например, первая предпослыка говорит о том, что мы должны были верно специфицировать модель. А что, если это не так и мы пропустили какую-нибудь переменную?

__[б]__ Оцените уравнение 

$$
TEST_i = \beta_0 + \beta_1 \cdot CLASS_i + \beta_2 \cdot EXPN_i + \varepsilon_i
$$

<!-- grade_id: cell-67ad004938dfa0b5, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

- Значимо ли уравнение в целом?  Запишите получившееся значение $F$-статистики для проверки этой гипотезы в переменную `f_stat`

<!-- grade_id: cell-2609963b6fb1daf8, points: 0, grade: false, solution: true -->
```python
f_stat = ...

# your code here

```

<!-- grade_id: cell-c9e8ea9ef25f5580, points: 1, grade: true, solution: false -->
```python
assert f_stat > 8

# несколько похожих скрытых тестов
```

Коэффициент перед переменной `CLASS` значимый, но он увеличился. По всей видимости, переменная `EXPN`, которую мы добавили, значимо влияет на результаты школьника. Кроме того, она коррелирует с `CLASS`. Когда мы не включаем её в модель, влияние этой переменной остаётся в случайной ошибке. Из-за этого возникает ковариация $Cov(CLASS_i, \varepsilon_i)$. 

Такая проблема называется __эндогенностью.__ Она приводит к смещению и несостоятельности оценок коэффициентов. Когда мы добавили в модель переменную `EXPN`, мы устранили возникшее у нас смещение. 

__[в]__ А что, если у нас есть какие-то другие переменные, которые оказались в ошибке, но при этом коррелируют с `CLASS`? Оцените уравнение:

$$
TEST_i = \beta_0 + \beta_1 \cdot CLASS_i + \beta_2 \cdot EXPN_i + \beta_3 \cdot INCOME + \varepsilon_i
$$

<!-- grade_id: cell-6469ad10816d0fe2, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

Обратите внимание, что значение перед переменной `CLASS` снова изменилось.

Изначально нас волновал вопрос: __Влияет ли размер класса на успеваемость?__ По факту, нас интересует только коэффициент перед этой переменной. Чтобы получить для него хорошую оценку, нам нужно учитывать в нашей регрессии все существенные факторы (т.е. все коррелированные с интересущей нас переменной факторы, коэффициенты при которых в истинной модели регрессии отличны от нуля). Такие дополнительные переменные обычно называют __контрольными__. 

Иногда бывает так, что интересующая нас переменная связана с чем-то ненаблюдаемым. Например, в данной ситуации успешность людей может быть связана с их талантом. Если мы никак не учтем талант людей в модели, мы получим неверные выводы. В эконометрике существуют различные продвинутые методы, которые позволяют попытаться оценить модели, учитывающие такие факторы. 

Узнать о них подробнее можно в [курсе по Экнонометрике,](https://www.coursera.org/learn/ekonometrika) В качестве учебника можно почитать, например, книгу "Введение в эконометрику" Филиппа Картаева. :)  