<center>
<img src="logo.png" height="900"> 
</center>


# Байесовская домашка

В этом задании мы оценим несколько простых моделей с помощью `pymc3`. 

Задачи делались [на основе книги Байесовские методы для хакеров.](https://github.com/CamDavidsonPilon/Probabilistic-Programming-and-Bayesian-Methods-for-Hackers) В ней вы можете найти ещё больше примеров использования байесовских методов с помощью pymc3. 

```python
import numpy as np
import pandas as pd

from scipy import stats as sts

import seaborn as sns
import matplotlib.pyplot as plt

plt.style.use('ggplot')
%matplotlib inline
```

```python
import warnings 
warnings.filterwarnings('ignore')
```

```python
import pymc3 as pm
pm.__version__
```

# Упражнение 1 (Маша и медведи)

Помните задачку с вынюхиванием из лекции? Давайте начнём с того, что попробуем решить её в `pymc3`.

__Модель медведя:__ $y_i \sim N(m, 2^2)$

__Априорное мнение, где спряталась Маша:__ $m \sim N(1, 4^2)$

```python
y = np.array([0.5, -1]) # медведи
```

Задайте априорное распределение и нашу модель.

<!-- grade_id: cell-9ab0e391a7a77635, points: 0, grade: false, solution: true -->
```python
basic_model = pm.Model()

with basic_model:
    # your code here
    
```

Найдите моду апостериорного распределения (байесовские методы для бедных). Запишите результат в переменную `map_estimate`.

<!-- grade_id: cell-f45f9b5beb2a0490, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

map_estimate = ...

# your code here

```

<!-- grade_id: cell-eabea37db4c38f2a, points: 1, grade: true, solution: false -->
```python
assert np.abs(float(map_estimate['m']) + 0.111111) < 1e-4

# несколько похожих скрытых тестов
```

Обучите модельс с помощью MCMC, 500 итераций для обучения будет достаточно. 

<!-- grade_id: cell-31e9242f37fbbd31, points: 0, grade: false, solution: true -->
```python
with basic_model:
    trace = ...
    # your code here
    
```

Посмотрите на визуализацию Марковских цепей. Правда ли, что цепь сошлась?

<!-- grade_id: cell-129b4b5de2613d8c, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

Найдите вероятность $P(m < 0)$. Ответ запишите в переменную `ans_prob`.

<!-- grade_id: cell-0c61232fe43634ca, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

<!-- grade_id: cell-6dacb7b1bb2a3597, points: 1, grade: true, solution: false -->
```python
assert ans_prob > 0.5

# несколько похожих скрытых тестов
```

> __Важно:__ перед тем, как отправить тетрадку на проверку, закоментируйте все части кода, где делаются долгие вычисления. На грейдер наложено ограничение в работе в 30 секунд и он может не успеть отработать. Из-за этого задание специфицировано таким образом, что вам нужно получить ответы на вопросы и вбить их в переменные.

# Упражнение 2 (Челленджер)

Байесовский подход - это просто другой способ оценивать те же самые модели. Например, давайте попробуем оценить логистическую регрессию. 

$28$ января $1986$ года космический шатл «Челленджер» взорвался при взлёте. Семь астронавтов, находившихся на борту, погибли. В ходе расследования причин катастрофы основной версией была неполадка с резиновыми уплотнительными кольцами в соединении с ракетными ускорителями. Для $23$ предшествовавших катастрофе полётов «Челленджера» известны температура воздуха и появление повреждений хотя бы у одного из уплотнительных колец.

```python
df = pd.read_csv('challenger.csv', sep=',')
print(df.shape)
df.head()
```

```python
t = df.Temperature.values
y = df.Incident.values

plt.figure(figsize=(10, 5))
plt.scatter(t, y, s=75, alpha=0.5)
plt.yticks([0, 1])

plt.ylabel("Была ли поломка")
plt.xlabel("Температура колец");
```


Обучим Байесовскую версию логистичесской регресcии.  Модель буде выглядеть следующим образом: 

\begin{equation*}
\begin{aligned}
& \beta \sim N(0, 10^2) \\
& \alpha \sim N(0, 10^2) \\
& p \sim Bern(p) \\
& p = P(y = 1) = \frac{1}{1 + e^{\alpha + \beta \cdot t}}
\end{aligned}
\end{equation*}

<!-- grade_id: cell-a3f473069b12ff42, points: 0, grade: false, solution: true -->
```python
import theano.tensor as tt  # под капотом пакета работает theano 
                            # именно на нём прописываются разные сложные формулы

t = df.Temperature.values
y = df.Incident.values

with pm.Model() as model:
    
    beta = ...
    alpha = ...
    # your code here
    
    p = pm.Deterministic("p", 1.0/(1. + tt.exp(beta*t + alpha)))
    y_obs = ...
    # your code here
    
```

Найдите моду апостериорного распределения (байесовские методы для бедных). Запишите результат в переменную `map_estimate`.

<!-- grade_id: cell-b039a0e67c27d2d0, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

map_estimate = ...

# your code here

```

<!-- grade_id: cell-3c2493dca40dc14e, points: 1, grade: true, solution: false -->
```python
assert np.abs(float(map_estimate['beta']) - 0.1665) < 1e-2

# несколько похожих скрытых тестов
```

Попробуем обучить модель немного по-новому. В переменной `step` запишемм конкретный алгоритм, который будем использовать при обучении. В `pm.sapmple` укажем из какой точке надо стартовать алгоритму. В качестве стартовой точки укажем моду апостериорного распределения, полученную в предыдущем пункте. 

```python
with model:
    step = pm.Metropolis()
    trace = pm.sample(120000, step=step, start=map_estimate)
    burned_trace = trace[100000::10]  # оставляем каждое 10 наблюдение, чтобы соседние наблюдения были менее зависимы
```

Посмотрим, сошлись ли цепи.

```python
az.plot_trace(trace);
```

Какова вероятность того, что $\beta < 0.0001$? То есть, какова вероятность того, что между температурой и вероятностью поломки нет связи? Запишите овтет в переменную `prob_beta`.

<!-- grade_id: cell-a335fe45dd9784db, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

prob_beta = ...

# your code here

```

<!-- grade_id: cell-78250ae46e19bebf, points: 1, grade: true, solution: false -->
```python
assert prob_beta < 0.1

# несколько похожих скрытых тестов
```

Нарисуем Получившиеся результаты. Для каждой температуры найдем $95\%$ байесовский интервал для вероятности того, что катастрофа произойдёт.

```python
from scipy.stats.mstats import mquantiles

def logistic(x, beta, alpha=0):
    return 1.0 / (1.0 + np.exp(np.dot(beta, x) + alpha))

t_seq = np.linspace(t.min() - 5, t.max()+5, 50)[:, None]
p_seq = logistic(t_seq.T, burned_trace['beta'][:,None], burned_trace['alpha'][:,None])

mean_prob_t = p_seq.mean(axis=0)
qs = mquantiles(p_seq, [0.025, 0.975], axis=0)

plt.figure(figsize=(10,5))
plt.fill_between(t_seq[:, 0], *qs, alpha=0.7, color="#7A68A6")
plt.plot(t_seq[:, 0], qs[0], label="95% интервал", color="#7A68A6", alpha=0.7)
plt.plot(t_seq, mean_prob_t, lw=1, ls="--", color="k", label="среднее")
plt.legend(loc="lower left")
plt.scatter(t, y, color="k", s=50, alpha=0.5)

plt.xlabel("Температура, $t$")
plt.ylabel("Вероятность поломки");
```

В день катастрофы температура уплотнителя была $31$ градус по Фарингейту. Найдите апостерионое распределение  вероятности катастрофы. 

<!-- grade_id: cell-b2c52f3d50449786, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

Найдите $5\%$ квантиль, математическое ожидание и $95\%$ квантиль апостериорного распределения. Заишите их в соответствующие переменные. Что это означает в контексте прогноза катастрофы? 

<!-- grade_id: cell-80e98850cd3adfb3, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

q05 = ...
avg = ...
q95 = ...

# your code here

```

<!-- grade_id: cell-4b191f33961305ca, points: 1, grade: true, solution: false -->
```python
assert q05  > 0.8

# несколько похожих скрытых тестов
```

> __Важно:__ перед тем, как отправить тетрадку на проверку, закоментируйте все части кода, где делаются долгие вычисления. На грейдер наложено ограничение в работе в 30 секунд и он может не успеть отработать. Из-за этого задание специфицировано таким образом, что вам нужно получить ответы на вопросы и вбить их в переменные.